```python
from IPython.display import HTML
style = "<style>div.dsc { background-color: #fcf2f2;border-color: #dFb5b4; border-left: 5px solid #dfb5b4; padding: 0.5em;}</style>"
HTML(style)
```

Output:
```
<IPython.core.display.HTML object>
```

<div class="dsc">
Duncans comments will appear in red boxes
</div>

<div class="dsc">Some thoughts / questions for this first HW.:    
1. Drop the image classification part of the HW
2. Add some basic operations with different data types -- dict, string, list, etc.
2. Will need to edit parts of markdown cells to make specific to this course.
3. can students add cells to these HW and how will those be evaluated / ignored.
</div>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---


<div class="dsc"> I'll edit this cell before the class starts.</div>

<br>
# HW1: Setup, Prerequisites



## Course Policies

Here are some important course policies. These are also located at: TBD

**Collaboration Policy**

Data science is a collaborative activity. While you may talk with others about
the homework, we ask that you **write your solutions individually**. If you do
discuss the assignments with others please **include their names** at the top
of your solution.

## This assignment

This part goes over prerequisites to taking DS100.

* How to set up Jupyter on your own computer.
* How to check out and submit assignments for this class.
* Python basics, like defining functions.
* How to use the `numpy` library to compute with arrays of numbers.
* Partial derivatives and matrix expressions


## Due Date



## Part 1: Prerequisites

### Setup

If you haven't already, go through the instructions at
http://www.ds100.org/sp18/setup.

The instructions for submission are at the end of this notebook.

You should now be able to open this notebook in Jupyter and run cells.

### Running a Cell

Try running the following cell.  If you unfamiliar with Jupyter Notebooks consider skimming [this tutorial](http://nbviewer.jupyter.org/github/jupyter/notebook/blob/master/docs/source/examples/Notebook/Notebook%20Basics.ipynb) or selecting **Help -> User Interface Tour** in the menu above. 

```python
print("Hello World!")
```

Output:
```
Hello World!
```

Even if you are familiar with Jupyter, we strongly encourage you to become proficient with keyboard shortcuts (this will save you time in the future).  To learn about keyboard shortcuts go to **Help -> Keyboard Shortcuts** in the menu above. 

Here are a few we like:
1. `ctrl`+`return` : *Evaluate the current cell*
1. `shift`+`return`: *Evaluate the current cell and move to the next*
1. `esc` : *command mode* (required before using any of the commands below)
1. `a` : *create a cell above*
1. `b` : *create a cell below*
1. `d` : *delete a cell*
1. `m` : *convert a cell to markdown*
1. `y` : *convert a cell to code*

### Testing your Setup

If you've set up your environment properly, this cell should run without problems:

```python
import math
import numpy as np
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import pandas as pd
# probably don't need skimage stuff.
import skimage
import skimage.io
import skimage.filters
```

### Python

Python is the main programming language we'll use in the course. We expect that you've taken CS61A or an equivalent class, so you should be able to explain the following cells. Run them and make sure you understand what is happening in each.

If this seems difficult, please review one or more of the following materials.

- **[Python Tutorial](https://docs.python.org/3.5/tutorial/)**: Introduction to Python from the creators of Python
- **[Composing Programs Chapter 1](http://composingprograms.com/pages/11-getting-started.html)**: This is more of a introduction to programming with Python.
- **[Advanced Crash Course](http://cs231n.github.io/python-numpy-tutorial/)**: A fast crash course which assumes some programming background.


<img src="launch.png" width="40px" align="left" float="left"/>
#### Mathematical Expressions
Note that the rocket icon indicates that you should just run the following cells.

```python
# This is a comment.
# In Python, the ** operator performs exponentiation.
math.sqrt(math.e ** (-math.pi + 1))
```

Output:
```
0.3427354792736325
```

<img src="launch.png" width="40px" align="left" float="left"/>
#### Output and Printing

```python
"Why didn't this line print?"

print("Hello" + ",", "world!")

"Hello, cell" + " output!"
```

Output:
```
Hello, world!
'Hello, cell output!'
```

<div class="dsc">: give an explanation for why string above didn't print</div>

<img src="launch.png" width="40px" align="left" float="left"/>
#### For Loops

```python
# A for loop repeats a block of code once for each
# element in a given collection.
for i in range(5):
    if i % 2 == 0:
        print(2**i)
    else:
        print("Odd power of 2")
```

Output:
```
1
Odd power of 2
4
Odd power of 2
16
```

<img src="launch.png" width="40px" align="left" float="left"/>
#### List Comprehension

```python
[str(i) + " sheep." for i in range(1,5)] 
```

Output:
```
['1 sheep.', '2 sheep.', '3 sheep.', '4 sheep.']
```

```python
[i for i in range(10) if i % 2 == 0]
```

Output:
```
[0, 2, 4, 6, 8]
```

<img src="launch.png" width="40px" align="left" float="left"/>
#### Defining Functions

```python
def add2(x):
    """This docstring explains what this function does: it adds 2 to a number."""
    return x + 2
```

<img src="launch.png" width="40px" align="left" float="left"/>
#### Getting Help

```python
help(add2)
```

Output:
```
Help on function add2 in module __main__:

add2(x)
    This docstring explains what this function does: it adds 2 to a number.
```

You can close the window at the bottom by pressing `esc` several times. 

<img src="launch.png" width="40px" align="left" float="left"/>
#### Passing Functions as Values

```python
def makeAdder(amount):
    """Make a function that adds the given amount to a number."""
    def addAmount(x):
        return x + amount
    return addAmount

add3 = makeAdder(3)
add3(4)
```

Output:
```
7
```

```python
makeAdder(3)(4)
```

Output:
```
7
```

<img src="launch.png" width="40px" align="left" float="left"/>
#### Anonymous Functions and Lambdas

```python
# add4 is very similar to add2, but it's been created using a lambda expression.
add4 = lambda x: x + 4
add4(5)
```

Output:
```
9
```

<img src="launch.png" width="40px" align="left" float="left"/>
#### Recursion

```python
def fib(n):
    if n <= 1:
        return 1
    else:
        # Functions can call themselves recursively.
        return fib(n-1) + fib(n-2)

fib(6)
```

Output:
```
13
```

### Question 1

#### Question 1a
Write a function nums_reversed that takes in an integer `n` and returns a string
containing the numbers 1 through `n` including `n` in reverse order, separated
by spaces. For example:

    >>> nums_reversed(5)
    '5 4 3 2 1'

***Note:*** The ellipsis (`...`) indicates something you should fill in.  It *doesn't* necessarily imply you should replace it with only one line of code.


<img src="code.png" width="75px" align="left"/> 
<br/>
The code icon indicates that you should complete the following block of code.

<div class="dsc"> Note that you''ll need to get the formatting just right for this to work.  Also explain that the 'NotImplementedError' can be deleted</div>

<!-- grade_id: num-reversed, points: 0, grade: false, solution: true -->
```python
def nums_reversed(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Test your code in the cell below.

<!-- grade_id: num-reversed-tests, points: 1, grade: true, solution: false -->
```python
assert nums_reversed(5) == '5 4 3 2 1'
assert nums_reversed(1) == '1'

```

Output:
```
<ipython-input-17-db0c0b9d3f7b> in <module>()
----> 1 assert nums_reversed(5) == '5 4 3 2 1'
      2 assert nums_reversed(1) == '1'

<ipython-input-16-b3470430a2d2> in nums_reversed(n)
      1 def nums_reversed(n):
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

<div class="dsc">explain what assert will do and how you know your code is working.</div>

#### Question 1b

Write a function `string_splosion` that takes in a non-empty string like
`"Code"` and returns a long string containing every prefix of the input.
For example:

    >>> string_splosion('Code')
    'CCoCodCode'
    >>> string_splosion('data!')
    'ddadatdatadata!'
    >>> string_splosion('hi')
    'hhi'

**Hint:** Try to use recursion. Think about how you might answering the following two questions:
1. **[Base Case]** What is the `string_splosion` of the empty string?
1. **[Inductive Step]** If you had a `string_splosion` function for the first $n-1$ characters of your string how could you extend it to the $n^{th}$ character? For example, `string_splosion("Cod") = "CCoCod"` becomes `string_splosion("Code") = "CCoCodCode"`.

<img src="code.png" width="75px" align="left"/> 

<!-- grade_id: string-splosion, points: 0, grade: false, solution: true -->
```python
def string_splosion(string):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Test your code in the cell below.

<!-- grade_id: string-splosion-test, points: 1, grade: true, solution: false -->
```python
assert string_splosion('Code') == 'CCoCodCode'
assert string_splosion('fade') == 'ffafadfade'

```

#### Question 1c

Write a function `double100` that takes in a list of integers
and returns `True` only if the list has two `100`s next to each other.

    >>> double100([100, 2, 3, 100])
    False
    >>> double100([2, 3, 100, 100, 5])
    True

<img src="code.png" width="75px" align="left"/> 

<!-- grade_id: double100, points: 0, grade: false, solution: true -->
```python
def double100(nums):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: double100-tests, points: 1, grade: true, solution: false -->
```python
assert double100([3, 3, 100, 100]) == True
assert double100([5, 2, 5, 2]) == False
assert double100([4, 2, 4, 100, 100, 5]) == True
```

### NumPy and Tables

The `NumPy` library lets us do fast, simple computing with numbers in Python. 

You should be able to understand the code in the following cells. If not, review the following:

* [DS100 Numpy Review](http://ds100.org/fa17/assets/notebooks/numpy/Numpy_Review.html)
* [Condensed Numpy Review](http://cs231n.github.io/python-numpy-tutorial/#numpy)
* [The Official Numpy Tutorial](https://docs.scipy.org/doc/numpy-dev/user/quickstart.html)



**Jupyter pro-tip**: Pull up the docs for any function in Jupyter by running a cell with
the function name and a `?` at the end:

<img src="launch.png" width="40px" align="left" float="left"/>

```python
np.arange?
```

**Another Jupyter pro-tip**: Pull up the docs for any function in Jupyter by typing the function
name, then `<Shift>-<Tab>` on your keyboard. Super convenient when you forget the order
of the arguments to a function. You can press `<Tab>` multiple tabs to expand the docs.

Try it on the function below:

```python
np.linspace
```

You can use the tips above to help you deciper the following code.

<img src="launch.png" width="40px" align="left" float="left"/>

```python
# Let's take a 20-sided die...
NUM_FACES = 20

# ...and roll it 4 times
rolls = 4

# What's the probability that all 4 rolls are different? It's:
# 20/20 * 19/20 * 18/20 * 17/20
prob_diff = np.prod((NUM_FACES - np.arange(rolls))
                    / NUM_FACES)
prob_diff
```

```python
# Let's compute that probability for 1 roll, 2 rolls, ..., 20 rolls.
# The array ys will contain:
# 
# 20/20
# 20/20 * 19/20
# 20/20 * 18/20
# ...
# 20/20 * 19/20 * ... * 1/20

xs = np.arange(20)
ys = np.cumprod((NUM_FACES - xs) / NUM_FACES)

# Python slicing works on arrays too
ys[:5]
```

```python
plt.plot(xs, ys, 'o-')
plt.xlabel("Num Rolls")
plt.ylabel('P(all different)')
```

```python
# Mysterious...
mystery = np.exp(-xs ** 2 / (2 * NUM_FACES))
mystery
```

```python
# If you're curious, this is the exponential approximation for our probability:
# https://textbook.prob140.org/ch1/Exponential_Approximation.html
plt.plot(xs, ys, 'o-', label="All Different")
plt.plot(xs, mystery, 'o-', label="Mystery")
plt.xlabel("Num Rolls")
plt.ylabel('P(all different)')
plt.legend()
```

### Question 2

To test your understanding of Numpy we will work through some basic image exercises.  In the process we will explore  visual perception and color.

Images are 2-dimensional grids of pixels. Each pixel contains 3 values between 0 and 1 that specify how much red, green, and blue go into each pixel.

We can create images in NumPy:

<img src="launch.png" width="40px" align="left" float="left"/>

<div class="dsc">I'd like another example for numpy -- this one not well explained.</div>

```python
simple_image = np.array([
    [[  0,   0, 0], [0.5, 0.5, 0.5], [1.0, 1.0, 1.0]], # Grayscale pixels
    [[1.0,   0, 0], [  0, 1.0,   0], [  0,   0, 1.0]], # Pure RGB pixels
    [[0.5, 0.5, 0], [0.5,   0, 0.5], [  0, 0.5, 0.5]], # Blend of 2 colors
])
simple_image
```

We can then use the [`scikit-image`](http://scikit-image.org/) library to display an image:

<img src="launch.png" width="40px" align="left" float="left"/>

```python
# Curious how this method works? Try using skimage.io.imshow? to find out.
# Or, you can always look at the docs for the method.
skimage.io.imshow(simple_image)
plt.grid(False) # Disable matplotlib's grid lines
```

We can read in image files using the `skimage.io.imread` method.  

*Note that in many image formats (e.g., JPEG) image values are numbers between 0 and 255 corresponding to a byte.  Therefore we divide each pixel value by 255 to obtain numbers between 0 and 1.*  

<img src="launch.png" width="40px" align="left" float="left"/>

```python
plt.figure(figsize=(20,10))

# Some image files (including .jpg files) have pixel values in between
# 0 and 255 when read. We divide by 255 to scale the values between 0 and 1:
pic = skimage.io.imread('target.jpg')/255


skimage.io.imshow(pic)
plt.grid(False) # Disable matplotlib's grid lines
```

*Professor Gonzalez is a **very amateur** archer*.

#### Question 2a

Complete the following block of code to plot the Red, Green, and Blue color channels separately.  The resulting images should appear in black and white.

* **Hint:** `pic[:, :, 0]` will slice the image to extract the red color channel.  Plotting the resulting matrix will generate a black and white picture.*

<img src="code.png" width="75px" align="left"/> 


```python
pic[:,:,0]
```

<!-- grade_id: imshowing-you-the-answer, points: 0, grade: false, solution: true -->
```python
plt.figure(figsize=(20,10)) 
channel_names = ["Red", "Green", "Blue"]

# Loop through index of each channel
for channel in range(3):
    # Make a subplot
    plt.subplot(1,3,channel+1)
    # FINISH THE CODE
    # Hint: you should write one line of code that calls
    # `skimage.io.imshow`
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    plt.grid(False)
    plt.title(channel_names[channel])
```

#### Question 2b

Surprisingly the human eye doesn't see all colors equally.  To demonstrate this we will study how blurring color channels affects image appearance.  First, we will try to blur each color channel individually.  Complete the following block of code using the `skimage.filters.gaussian` blurring function (read the docs) to render a blurry black and white image for each of the color channels.  You should set the standard deviation of the Gaussian blurring kernel `sigma` to 10.

<img src="code.png" width="75px" align="left"/> 

<!-- grade_id: imshow-blur, points: 0, grade: false, solution: true -->
```python
plt.figure(figsize=(20,10))

sigma = 10

# Loop through index of each channel
for channel in range(3):
    # Make a subplot
    plt.subplot(1,3,channel+1)
    # FINISH THE CODE 
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    plt.grid(False)
    plt.title(channel_names[channel])
```

#### Question 2c

Using the following block of code:

```python
pic_copy = pic.copy()
pic_copy[:, :, channel] = ...
skimage.io.imshow(pic_copy)
```

we can replace a color channel with a different black and white image.  Complete the following block of code to render three different versions of the full color image with just one of the channels blurred.

<img src="code.png" width="75px" align="left"/> 

<!-- grade_id: blur-the-channel, points: 0, grade: false, solution: true -->
```python
plt.figure(figsize=(20,10))

sigma = 10

# Loop through index of each channel
for channel in range(3):
    # Make a subplot
    plt.subplot(1,3,channel+1)
    # FINISH THE CODE
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    plt.grid(False)
    plt.title(channel_names[channel])
```

#### Question 2d

Each image should appear slightly different.  Which one is the blurriest and which is the sharpest?  Write a short description of what you see in the cell below.


<img src="write.png" width="120px" align="left" float="left"/> *This icon means you will need to write in text response in the cell below using English.

***Hint:** I observe ... .  On possible explanation for this is ... .


<!-- grade_id: question2d-answer, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<div class="dsc">
## Possible new section:
Have students pull data from the CAISO "renewables watch" API and put all data into a numpy matrix.  
1. Need to use python to pull data from web.  E.g. http://content.caiso.com/green/renewrpt/20180617_DailyRenewablesWatch.txt.  Replace 20180617 with YYYYMMDD to get other days.
2. Massage text files into numpy matrix. 
3. Then merge multiple days together
4. Can answer a question like: "on what day was solar power production the highest in 2017"? </div>

<div class="dsc">For question 3, explain to students that they just have to try their best -- will help to understand how to deliver content.</div>

## Multivariable Calculus and Linear Algebra

The following questions ask you to recall your knowledge of multivariable calculus and linear algebra. We will use some of the most fundamental concepts from each discipline in this class, so the following problems should at least seem familiar to you.

For the following problems, you should use LaTeX to format your answer. If you aren't familiar with LaTeX, not to worry. It's not hard to use in a Jupyter notebook. Just place your math in between dollar signs:

\$ f(x) = 2x \$ becomes $ f(x) = 2x $.

If you have a longer equation, use double dollar signs:

\$\$ \sum_{i=0}^n i^2 \$\$ becomes:

$$ \sum_{i=0}^n i^2 $$.

Here are some handy notation:

| Output | Latex   |
|:--|:--|
| $$x^{a + b}$$  | `x^{a + b}` |
| $$x_{a + b}$$ | `x_{a + b}` |
| $$\frac{a}{b}$$ | `\frac{a}{b}` |
| $$\sqrt{a + b}$$ | `\sqrt{a + b}` |
| $$\{ \alpha, \beta, \gamma, \pi, \mu, \sigma^2  \}$$ | `\{ \alpha, \beta, \gamma, \pi, \mu, \sigma^2  \}` |
| $$\sum_{x=1}^{100}$$ | `\sum_{x=1}^{100}` |
| $$\frac{\partial}{\partial x} $$ | `\frac{\partial}{\partial x} ` |
| $$\begin{bmatrix} 2x + 4y \\ 4x + 6y^2 \\ \end{bmatrix}$$ | `\begin{bmatrix} 2x + 4y \\ 4x + 6y^2 \\ \end{bmatrix}`|



[For more about basic LaTeX formatting, you can read this article.](https://www.sharelatex.com/learn/Mathematical_expressions)

If you have trouble with these topics, we suggest reviewing:

- [Khan Academy's Multivariable Calculus](https://www.khanacademy.org/math/multivariable-calculus)
- [Khan Academy's Linear Algebra](https://www.khanacademy.org/math/linear-algebra)

### Question 3

#### Question 3a

Simplify the following expression:

$$
\large \ln \left( 3 e^{2  x} e^{\frac{1}{x^2}} \right)
$$

<img src="write.png" width="120px" align="left" float="left"/> *This icon means you will need to write in text response in the cell below using English + $\LaTeX$.

$$
\ln \left( 3 e^{2  x} e^{\frac{1}{x^2}} \right) = 
\ldots
$$

<!-- grade_id: question3a-answer, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

#### Question 3b


Suppose we have the following scalar-valued function on $x$ and $y$:

$$ \Large f(x, y) = x^2 + 4xy + 2y^3 + e^{-3y} + \ln(2y) $$

Compute the partial derivative with respect to $x$.

<img src="write.png" width="75px" align="left" float="left"/>

$$ \frac{\partial}{\partial x} f(x,y) = ... $$

<!-- grade_id: question3b-answer, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

#### Question 3c


Now compute the partial derivative of $f(x,y)$ with respect to $y$:

<img src="write.png" width="75px" align="left" float="left"/>

$$ \frac{\partial}{\partial y} f(x,y) = ... $$

<!-- grade_id: question3c-answer, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Question 4

In this question, we'll ask you to use your linear algebra knowledge to fill in NumPy matrices. To conduct matrix multiplication in NumPy, you should write code like the following:

<img src="launch.png" width="40px" align="left" float="left"/>

```python
# A matrix in NumPy is simply a 2-dimensional NumPy array
matA = np.array([
    [1, 2, 3],
    [4, 5, 6],
])

matB = np.array([
    [10, 11],
    [12, 13],
    [14, 15],
])

# The notation B @ v means: compute the matrix multiplication Bv
matA @ matB
```

You can also use the same syntax to do matrix-vector multiplication or vector dot products. Handy!

<img src="launch.png" width="40px" align="left" float="left"/>

```python
matA = np.array([
    [1, 2, 3],
    [4, 5, 6],
])

# A vector in NumPy is simply a 1-dimensional NumPy array
some_vec = np.array([ 10, 12, 14, ])

another_vec = np.array([ 10, 20, 30 ])

print(matA @ some_vec)
print(some_vec @ another_vec)
```

#### Question 4a

Joey, Deb, and Sam are shopping for fruit at Berkeley Bowl. Berkeley Bowl, true to its name, only sells fruit bowls. A fruit bowl contains some fruit and the price of a fruit bowl is the total price of all of its individual fruit.

Berkeley Bowl has apples for \$2.00, bananas \$1.00, and cantaloupes \$4.00 (expensive!). The price of each of these can be written in a vector:

$$
\vec{v} = \begin{bmatrix}
     2 \\
     1 \\
     4 \\
\end{bmatrix}
$$

Berkeley Bowl sells the following fruit bowls:

1. 2 of each fruit
2. 5 apples and 8 bananas
3. 2 bananas and 3 cantaloupes
4. 10 cantaloupes

Create a 2-dimensional numpy array encoding the matrix $B$ such that the matrix-vector multiplication

$$
B\vec{v}
$$

evaluates to a length 4 column vector containing the price of each fruit bowl. The first entry of the result should be the cost of fruit bowl #1, the second entry the cost of fruit bowl #2, etc.

<img src="code.png" width="75px" align="left" float="left"/>

<!-- grade_id: question4-code, points: 0, grade: false, solution: true -->
```python
v = np.array([2,1,4])

# Fill in B
# B = ...

# YOUR CODE HERE
raise NotImplementedError()

# The notation B @ v means: compute the matrix multiplication Bv
B @ v
```

<!-- grade_id: question4-tests, points: 1, grade: true, solution: false -->
```python
assert np.allclose(B @ v, np.array([14, 18, 14, 40]))
```

#### Question 4b

Joey, Deb, and Sam make the following purchases:

- Joey buys 2 fruit bowl #1's and 1 fruit bowl #2.
- Deb buys 1 of each fruit bowl.
- Sam buys 10 fruit bowl #4s (he really like cantaloupes).

Create a matrix $A$ such that the matrix expression

$$
AB\vec{v}
$$

evaluates to a length 3 column vector containing how much each of them spent. The first entry of the result should be the total amount spent by Joey, the second entry the amount sent by Deb, etc.

<img src="code.png" width="75px" align="left" float="left"/>

<!-- grade_id: question4b-code, points: 0, grade: false, solution: true -->
```python
A = np.array([
    [2, 1, 0, 0],
    # Finish this!
    
    # YOUR CODE HERE
    raise NotImplementedError()
]) 

A @ B @ v 
```

<!-- grade_id: question4b-test, points: 1, grade: true, solution: false -->
```python
# The final tests for this question have been hidden.

```

#### Question 4c

Who spent the most money?

<img src="write.png" width="120px" align="left" float="left"/> 

<!-- grade_id: question4c-answer, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

#### Question 4d

Let's suppose Berkeley Bowl changes their fruit prices, but you don't know what they changed their prices to. Joey, Deb, and Sam buy the same quantity of fruit baskets and the number of fruit in each basket is the same, but now they each spent these amounts:

$$
\vec{x} = \begin{bmatrix}
    80 \\
    80 \\
    100 \\
\end{bmatrix}
$$

Use `np.linalg.inv` and the above final costs to compute the new prices for the individual fruits:

<img src="code.png" width="75px" align="left" float="left"/> 

<!-- grade_id: question4d, points: 0, grade: false, solution: true -->
```python
# Fill in
new_v = ...

# YOUR CODE HERE
raise NotImplementedError()

new_v
```

<!-- grade_id: question4d-test, points: 1, grade: true, solution: false -->
```python
assert np.allclose(new_v, np.array([ 5.5,  2.20833333,  1.]))
```

# Submission

You're done!

In order to turn in this assignment, submit this notebook to the Data 100 datahub at http://data100.datahub.berkeley.edu. 

You will need to upload this notebook and any associated files to datahub manually if you have completed this assignment on your local machine. Detailed instructions for how to submit on Datahub can be found at http://www.ds100.org/sp18/materials.

Remember to click 'Validate' for this assignment before submitting. After clicking 'Submit', you can verify there is a time-stamped submission under 'Submitted assignments'.